```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment3.ipynb")
```

```python
import numpy as np
import matplotlib.pyplot as plt
```

# Probability, Statistics and Distances Homework

Please complete the following assignment. Run all cells and submit your completed notebook through Gradescope.

This assignment covers the following topics:

- Frequentist vs Bayesian probability
- Definition of probability
- Conditional probability and Bayes theorem
- Random variables
- Expectation, variance, and covariance
- Binomial distribution
- Poisson distribution
- Gaussian distribution
- Jaccard Index
- Euclidean Distances

<!-- BEGIN QUESTION -->

#### 1. Which of the following statements best describes the Bayesian interpretation of probability?

    A) Probability is the long-run relative frequency of an event occurring in repeated trials.
    B) Probability represents a degree of belief or confidence in the occurrence of an event.
    C) Probability is the ratio of favorable outcomes to the total number of equally likely outcomes.
    D) Probability is the measure of the likelihood that a random variable takes a specific value.

    Write a function `ans1` and return your answer as a string, e.g., "A".

```python
def ans1():
    return "B"
```

<!-- END QUESTION -->

#### 2. Write a function `roll_die(n:int) -> list` that simulates rolling a six-sided die `n` times and returns a `list` of outcomes.

```python
import numpy as np
def roll_die(n: int) -> list:
    return np.random.default_rng().integers(low = 1, high = 6, size = n)
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

#### 3. Suppose that 1% of a population has a particular disease. A test for the disease correctly identifies  diseased individuals 99% of the time (sensitivity), and correctly identifies non-diseased individuals  95% of the time (specificity). If a randomly selected individual tests positive, what is the probability that they actually have the disease?
    Return your answer as a float between 0 and 1.

```python
def answer3():
    P_D = 0.01
    P_not_D = 0.99
    P_T_pos_given_D = 0.99
    P_T_pos_given_not_D = 1 - 0.95
    ans = (P_D * P_T_pos_given_D) / (P_D * P_T_pos_given_D + P_not_D *  P_T_pos_given_not_D)
    return ans
```

```python
answer3()
```

Output:
```
0.16666666666666655
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

#### Let X be a discrete random variable that takes the values 1, 2, 3 with probabilities P(X=1)=0.2, P(X=2)=0.5, P(X=3)=0.3. 
-   Compute the expected value E[X].
-   Return your answer as a float.

```python
def ans4():
    return 1*0.2+2*0.5+3*0.3
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

####  5. Let X be a random variable as in Question 4. Compute the variance Var(X).

-    Return your answer as a float.

```python
def ans5():
    E_X = ans4()
    E_X2 = 1*0.2+2**2*0.5+3**2*0.3
    return  E_X2 - E_X**2
```

```python
ans5()
```

Output:
```
0.490000000000002
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

#### 6. A fair coin is tossed 10 times. What is the probability of getting exactly 6 heads?
    Return your answer as a float.

```python
from math import comb
def ans6():
    return 10*9*8*7/(4*3*2*1)*0.5**10
```

```python
ans6()
```

Output:
```
0.205078125
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

#### 7. A radioactive source emits particles following a Poisson distribution with an average rate of λ = 4 particles per minute. 
    What is the probability that exactly 6 particles are emitted in a given minute?

    Return your answer as a float.

```python
from scipy.stats import poisson
def ans7():
    λ = 4
    k = 6
    return float(poisson.pmf(k, λ))
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

#### 8. Assuming that the heights of adult males are normally distributed with a mean of 70 inches and a standard deviation of 3 inches, what is the probability that a randomly selected male is taller than 75 inches?

    Return your answer as a float.

```python
from scipy.stats import norm
def ans8():
    mu = 70
    sigma = 3
    x = 75
    return 1 - float(norm(mu, sigma).cdf(x))
```

```python
ans8()
```

Output:
```
0.047790352272814696
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

#### 9. Write a function to calculate the Jaccard index (also known as Jaccard similarity coefficient) between two sets.

    Example:
    Input: set1 = {1, 2, 3, 4, 5}, set2 = {2, 4, 6, 8}
    Output: 0.2857
    Explanation:
    Intersection: {2, 4}
    Union: {1, 2, 3, 4, 5, 6, 8}
    Jaccard Index = 2 / 7 ≈ 0.2857

```python
def ans9(set1: set, set2: set) -> float:
    intersec = set1.intersection(set2)
    union = set1.union(set2)
    return float(len(intersec)) / float(len(union))
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

#### 10. Write a function that computes the norms between two lists.  (1-norm, 2-norm, infinity-norm, euclidean distance)
    Return a dictionary with keys `1_norm`, `2_norm`, `inf_norm` and `euclidean`.
    Note: Convert all the norms to float while returning.

```python
def ans10(arr1: np.array, arr2: np.array) -> float:
    dict_norm = {'1_norm': float(np.linalg.norm(arr1 - arr2, ord = 1)), '2_norm': float(np.linalg.norm(arr1 - arr2, ord = 2)), 'inf_norm': float(np.linalg.norm(arr1 - arr2, ord = np.inf)), 'euclidean': float(np.linalg.norm(arr1 - arr2, ord = 2))}
    return dict_norm
```

```python
grader.check("q10")
```

Output:
```
q10 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q10 (points: None)
```
>>> assert type(ans10(np.array([1, 2, 3]), np.array([4, 5, 6]))) == dict
>>> assert ans10(np.array([1, 2, 3]), np.array([4, 5, 6]))['1_norm'] == 9.0
>>> assert ans10(np.array([1, 2, 3]), np.array([4, 5, 6]))['inf_norm'] == 3.0
>>> assert ans10(np.array([1, 2, 3]), np.array([4, 5, 6]))['2_norm'] > 5.0
```


Tests for q2 (points: None)
```
>>> assert len(roll_die(10)) == 10
>>> assert all((1 <= x <= 6 for x in roll_die(100)))
```


Tests for q3 (points: None)
```
>>> assert answer3() > 0.15 and answer3() < 0.2
>>> assert type(answer3()) == float
```


Tests for q4 (points: None)
```
>>> assert type(ans4()) == float
>>> assert 2 < ans4() < 2.1
```


Tests for q5 (points: None)
```
>>> assert type(ans5() == float)
>>> assert abs(ans5() - 0.49) < 1e-09
>>> assert 0.49 < ans5() < 0.491
```


Tests for q6 (points: None)
```
>>> assert type(ans6()) == float
>>> assert 0.2 < ans6() < 0.21
>>> assert abs(ans6()) - 0.2 < 0.05
```


Tests for q7 (points: None)
```
>>> assert type(ans7()) == float
>>> assert 0.1 < ans7() < 0.11
>>> assert abs(ans7()) - 0.1 < 0.0042
```


Tests for q8 (points: None)
```
>>> assert type(ans8()) == float
>>> assert 0.04 < abs(ans8()) < 0.05
```


Tests for q9 (points: None)
```
>>> assert ans9({1, 2, 3, 4, 5}, {2, 3, 4}) == 0.6
>>> assert 0.4 < ans9({1, 2, 3, 4, 5, 6, 7}, {2, 3, 4}) < 0.5
>>> assert 0.2 < ans9({1, 2, 3, 4, 5, 6, 7, 8, 9}, {2, 3}) < 0.23
```
